# Activity: Using Lists & Tuples

## Introduction

In this activity you will practice using and creating lists and tuples.
This activity contains the following topics:
- Creating lists
- Creating tuples
- Indexing
- Slicing
- List methods
- Unpacking


#### Question 1

Create a list called `countries` with the names of 5 different countries of your choice.


<!-- grade_id: cell-0614fe54fc242fa2, points: 0, grade: false, solution: true -->
```python
# Your code here
countries = ['philippines', 'thailand', 'indonesia', 'malaysia', 'singapore']
```

<!-- grade_id: cell-3a329dfd53eab709, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(countries, list), 'Your variable is not a list'

```

#### Question 2

Using the `countries` list, assign the first and last countries in the list to the variables `first_country` and `last_country` respectively.


<!-- grade_id: cell-2d238064c27e485b, points: 0, grade: false, solution: true -->
```python
# Your code here
countries = ['philippines', 'thailand', 'indonesia', 'malaysia', 'singapore']
first_country = 'philippines'
last_country = 'singapore'
print(first_country)
print(last_country)
```

Output:
```
philippines
singapore
```

<!-- grade_id: cell-10e571d348bfb921, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert len(countries) == 5, 'Your list does not have 5 elements'
assert isinstance(first_country, str), 'Your variable is not a string'
assert isinstance(last_country, str), 'Your variable is not a string'

```

#### Question 3

Using the `countries` list, add a country to the end of the list.
Then assign the last 3 countries to a variable called `last_three_countries`.


<!-- grade_id: cell-09f04a4add4b7637, points: 0, grade: false, solution: true -->
```python
# Your code here
countries = ['philippines', 'thailand', 'indonesia', 'malaysia', 'singapore']
countries.append('vietnam')
last_three_countries = ['malaysia', 'singapore', 'vietnam']
print(last_three_countries)
len(last_three_countries)



```

Output:
```
['malaysia', 'singapore', 'vietnam']
3
```

<!-- grade_id: cell-2067ef95fa07586c, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert len(countries) == 6, 'Your list does not have 6 elements (one more element on the end)'
assert isinstance(last_three_countries, list), 'Your variable is not a list'
assert len(last_three_countries) == 3, 'Your list does not have 3 elements'

```

#### Question 4

Using the `countries` list, assign the **even** indexed countries to a variable called `even_countries`. Remember the **first index is 0** (an even number).

This means if your list was something like:
```python
countries = [
    'Lilliput',
    'Gilead',
    'Oceania',
    'Wakanda',
    'Genovia',
    'Panem',
]
```
Then `even_countries` will be:
```python
[
    'Lilliput',
    'Oceania',
    'Genovia',
]
```


<!-- grade_id: cell-4a5264eeab7a9c33, points: 0, grade: false, solution: true -->
```python
# Your code here
countries = ['philippines', 'thailand', 'indonesia', 'malaysia', 'singapore']
countries.append('vietnam')
even_countries = ['philippines', 'indonesia', 'singapore']
print(even_countries)
len(even_countries)







```

Output:
```
['philippines', 'indonesia', 'singapore']
3
```

<!-- grade_id: cell-e753d72d2a6b5400, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(even_countries, list), 'Your variable is not a list'

```

#### Question 5

Add another country to the `countries` list and assign the length of the list to a variable called `num_countries`.


<!-- grade_id: cell-8162b18aba873a5b, points: 0, grade: false, solution: true -->
```python
# Your code here
countries = ['philippines', 'thailand', 'indonesia', 'malaysia', 'singapore']
countries.append('vietnam')
countries.append('cambodia')
num_countries = len(countries) 
print(num_countries)





```

Output:
```
7
```

<!-- grade_id: cell-5d6cf2f545b99477, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert isinstance(num_countries, int), 'Your variable is not an integer'
assert num_countries > 6, 'Check the number of countries; you should have only added 2 more countries after the initial 5'

```

#### Question 6

You're given a tuple containing superheroes, unpack the tuple into variables `superman`, `batman`, and `black_panther`
 (in the expected order).

```python
# Given tuple

superheroes = ('Superman', 'Batman', 'Black Panther')
```

<!-- grade_id: cell-a6d6bddc63d43806, points: 0, grade: false, solution: true -->
```python
# Your code here
superheroes = ('Superman', 'Batman', 'Black Panther')
superman, batman, black_panther = superheroes
print(superman)
print(batman)
print(black_panther)

```

Output:
```
Superman
Batman
Black Panther
```

<!-- grade_id: cell-adfc0a04f765e223, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert isinstance(superman, str), 'The superman variable is not a string'
assert isinstance(batman, str), 'The batman variable is not a string'
assert isinstance(black_panther, str), 'The black_panther variable is not a string'

```

#### Question 7

Create a list of **5 tuples** called `planet_info` where each tuple contains a planet name, the planet's color, and the planet's number of moons.


For example, your list might look something like this:
```python
planet_info = [
    ('Coruscant', 'orange', 4),
    ...
    ('Krypton', 'green', 3),
]
```

<!-- grade_id: cell-5275e17599ccffcf, points: 0, grade: false, solution: true -->
```python
# Your code here
planet_info = [
    ('Coruscant', 'orange', 4),
    ('Coco', 'white', 5),
    ('Kiki', 'blue', 0),
    ('Hatdog', 'red', 3),
    ('Krypton', 'green', 3),
]
planet_info


```

Output:
```
[('Coruscant', 'orange', 4),
 ('Coco', 'white', 5),
 ('Kiki', 'blue', 0),
 ('Hatdog', 'red', 3),
 ('Krypton', 'green', 3)]
```

<!-- grade_id: cell-ac994f5244998529, points: 1, grade: true, solution: false -->
```python
# Question 7 Grading Checks

assert isinstance(planet_info, list), 'Your variable is not a list'
assert isinstance(planet_info[0], tuple), 'Your first item is not a tuple'
assert len(planet_info[0]) == 3, 'Your first tuple does not have 3 elements'
assert len(planet_info) == 5, 'Your list does not have 5 elements'

```